# Introduction

A feature of modern programming languages is an extensive library of *standard functions*. This means that we can make use of well-tested and optimised functions for performing common tasks rather than writing our own. This makes our programs shorter and of higher quality, and in most cases faster.

## Objectives

- Introduce use of standard library functions
- Importing and using modules
- Introduction to namespaces
- Print formatting of floats
- Accessing documentation for functions and modules

# The standard library

You have already used some standard library types and functions. In previous activities we have used built-in types like `string` and `float`, and the function `abs` for absolute value. We have made use of the standard library function `print` to display to the screen.

Python has a large standard library. To organise it, most functionality is arranged into 'modules', with each module providing a range of related functions. Before you program a function, check if there is a library function that can perform the task. The Python standard library is documented at https://docs.python.org/3/library/.
Search engines are a good way to find library functions, e.g. entering "Is there a Python function to compute the hyperbolic tangent of a complex number" into a search engine will take you to the function `cmath.tanh`. Try this link: https://bfy.tw/7aMc.

```python

```

# Other libraries

The standard library tools are general purpose and will be available in any Python environment.
Specialised tools are usually made available in other libraries (modules). There is a huge range of Python libraries available for specialised problems. We have already used some parts
of NumPy (https://www.numpy.org/), which is a specialised library for numerical computation.

The simplest way to install a non-standard library is using the command `pip`. From the command line, the library NumPy is installed using:

    pip install numpy
    
and from inside a Jupyter notebook use:

    %pip install numpy

NumPy is so commonly used it is probably already installed on computers you will be using.
You will see `pip` being used in some later notebooks to install special-purpose modules.

When developing programs outside of learning exercises,
if there is a no standard library module for a problem you are trying to solve, 
search online for a module before implementing your own.

Let's make sure the libraries we need now and later are installed.

```python
%pip install numpy
%pip install matplotlib

```

Output:
```
Requirement already satisfied: numpy in /home/studio-lab-user/.conda/envs/default/lib/python3.9/site-packages (1.22.2)
Note: you may need to restart the kernel to use updated packages.
Requirement already satisfied: matplotlib in /home/studio-lab-user/.conda/envs/default/lib/python3.9/site-packages (3.5.1)
Requirement already satisfied: packaging>=20.0 in /home/studio-lab-user/.conda/envs/default/lib/python3.9/site-packages (from matplotlib) (21.3)
Requirement already satisfied: fonttools>=4.22.0 in /home/studio-lab-user/.conda/envs/default/lib/python3.9/site-packages (from matplotlib) (4.29.1)
Requirement already satisfied: python-dateutil>=2.7 in /home/studio-lab-user/.conda/envs/default/lib/python3.9/site-packages (from matplotlib) (2.8.2)
Requirement already satisfied: cycler>=0.10 in /home/studio-lab-user/.conda/envs/default/lib/python3.9/site-packages (from matplotlib) (0.11.0)
Requirement already satisfied: kiwisolver>=1.0.1 in /home/studio-lab-user/.conda/envs/default/lib/pytho
```

# Using library functions: `math` example

To use a function from a module we need to make it available in our program. 
This is called 'importing'. We have done this in previous notebooks with the `math` module, but without explanation. The process is explained below.

The `math` module (https://docs.python.org/3/library/math.html) provides a wide range of mathematical functions. For example, to compute the square root of a number, we do:

```python
import math

x = 2.0
x = math.sqrt(x)
print(x)
```

Output:
```
1.4142135623730951
```

Dissecting the above code block, the line 
```python
import math 
```
makes the math module available in our program. It is good style to put all `import` statements at the top of a file (or at the top of a cell when using a Jupyter notebook). 

The function call
```python
x = math.sqrt(x)
```    
says 'use the `sqrt` function from the `math` module to compute the square root'.

By prefixing `sqrt` with `math`, we are using a *namespace* (which in this case is `math`).
This makes clear precisely which `sqrt` function we want to use - there could be more than one `sqrt` function available.

> *Namespaces:* The prefix '`math`' indicates which '`sqrt`' function we want to use. This
might seem pedantic, but in practice there are often different algorithms for performing the same or similar  operation. They might vary in speed and accuracy. In some applications we might need an accurate (but slow)  method for computing the square root, while for other applications we might need speed with a compromise on accuracy. But, if two functions have the same name and are not distinguished by a name space, we have a *name clash*.

> In a large program, two developers might choose the same name for two functions that perform similar but slightly different tasks. If these functions are in different modules, there will be no name clash since the module name provides a 'namespace'  - a prefix that provides a distinction between the two functions. Namespaces are extremely helpful for multi-author programs. Older languages, like C and Fortran, might not support namespaces. Most modern languages do support namespaces.  

We can import specific functions from a module, e.g. importing only the `sqrt` function:

```python
from math import sqrt

x = 2.0
x = sqrt(x)
print(x)
```

Output:
```
1.4142135623730951
```

This way, we are importing (making available) only the `sqrt` function from the `math` module (the `math` module has a large number of functions).

We can even choose to re-name functions that we import:

```python
from math import sqrt as some_math_function 

x = 2.0
x = some_math_function(x)
print(x)
```

Output:
```
1.4142135623730951
```

Renaming functions at import can be helpful to keep code short, and we will see below it can be useful for switching between different functions.
However the above choice of name is very poor practice - the name '`some_math_function`' is not descriptive.
Below is a more sensible example.

Say we program a function that computes the roots of a quadratic function using the quadratic formula:

```python
from math import sqrt as square_root

def compute_roots(a, b, c):
    "Compute roots of the polynomial f(x) = ax^2 + bx + c"
    root0 = (-b + square_root(b*b - 4*a*c))/(2*a)
    root1 = (-b - square_root(b*b - 4*a*c))/(2*a)
    return root0, root1

# Compute roots of f = 4x^2 + 10x + 1
root0, root1 = compute_roots(4, 10, 1)
print(root0, root1)
```

Output:
```
-0.10435607626104004 -2.3956439237389597
```

The above is fine as long as the polynomial has real roots. However, the function `math.sqrt` 
will give an error (technically, it will 'raise an exception') if a negative argument is passed to it. This is to stop naive mistakes.

We do know about complex numbers, so we want to compute complex roots. The Python module `cmath` provides functions for complex numbers. If we were to use `cmath.sqrt` to compute the square root, our function would support complex roots. We do this by importing the `cmath.sqrt` functions as `square_root`:

```python
# Use the function from cmath as square_root to compute the square root
# (this will replace the previously imported sqrt function)
from cmath import sqrt as square_root

# Compute roots (roots will be complex in this case)
root0, root1 = compute_roots(40, 10, 1)
print(root0, root1)

# Compute roots (roots will be real in this case, but cmath.sqrt always returns a complex type)
root0, root1 = compute_roots(4, 10, 1)
print(root0, root1)
```

Output:
```
(-0.125+0.09682458365518543j) (-0.125-0.09682458365518543j)
(-0.10435607626104004+0j) (-2.3956439237389597+0j)
```

The function now works for all cases because `square_root` is now using `cmath.sqrt`. Note that `cmath.sqrt` always returns a complex number type, even when the complex part is zero.

# String functions and string formatting

A standard function that we have used from the start is '`print`'. This function turns arguments into a string and displays the string to the screen. So far, we have only printed simple variables and relied mostly on the default conversions to a string for printing to the screen (the exception was printing the floating point representation of 0.1, where we needed to specify the number of significant digits to see the inexact representation in binary).

## Formatting

We can control how strings are formatted and displayed. Below is an example of inserting a string variable and a number variable into a string of characters. Not the `f` before the string. This indicates it is a formatted string and that the variables should be inserted.

```python
# Format a string with name and age
name = "Amber"
age = 19
text_string = f"My name is {name} and I am {age} years old."

# Print to screen 
print(text_string)

# Short-cut for printing without assignment
name = "Ashley"
age = 21
print(f"My name is {name} and I am {age} years old.")
```

Output:
```
My name is Amber and I am 19 years old.
My name is Ashley and I am 21 years old.
```

For floating-point numbers we often want to control the formatting, and in particular the number of significant figures displayed. Using the display of $\pi$ as an example: 

```python
# Import math module to get access to math.pi
import math

# Default formatting
print(f"The value of π using the default formatting is: {math.pi}")

# Control number of significant figures in formatting
print(f"The value of π to 5 significant figures is: {math.pi:.5}")
print(f"The value of π to 8 significant figures is: {math.pi:.8}")
print(f"The value of π to 20 significant figures and using scientific notation is: {math.pi:.20e}")
```

Output:
```
The value of π using the default formatting is: 3.141592653589793
The value of π to 5 significant figures is: 3.1416
The value of π to 8 significant figures is: 3.1415927
The value of π to 20 significant figures and using scientific notation is: 3.14159265358979311600e+00
```

There are many more ways in which float formatting can be controlled - search online if you want to format a float in a particular way.  

# Module example: fractions and statistics

Python has standard library support for fractions and statistical operations. Summing fractions:

```python
from fractions import Fraction

f0 = Fraction(2, 3)
f1 = Fraction(3, 7)
result = f0 + f1
print("Fraction:", result)
print("Float:", float(result))
```

Output:
```
Fraction: 23/21
Float: 1.0952380952380953
```

We can use the `statistics` module perform statistical operations on a list of numbers. In this example we will represent the numbers using `fractions.Fraction`.

```python
import statistics

data = [Fraction(1, 2), Fraction(2, 1), Fraction(1, 3), Fraction(1, 17), Fraction(-1, 13)]
print("numbers:", data)

print("mean:", statistics.mean(data))
print("mode:", statistics.mode(data))
var = statistics.variance(data)
#print("variance: {} ({})".format(var, float(var))) #alterntiave way to format strings
print(f"variance : {var} ({float(var)})")
```

Output:
```
numbers: [Fraction(1, 2), Fraction(2, 1), Fraction(1, 3), Fraction(1, 17), Fraction(-1, 13)]
mean: 3733/6630
mode: 1/2
variance: 3060917/4395690 (0.6963450561800308)
variance : 3060917/4395690 (0.6963450561800308)
```

# Tab completion

Tab completion is useful for seeing which functions a module provides. For example, after importing the `math` module type `math.` and press `tab`. This will show all the available functions in `math`. Typing `math.c` and pressing `tab` will show all the available functions that start with 'c'. Why not try this below...

```python

```

# Accessing function documentation

The documentation for a function can be accessed using the `help` function, e.g.:

```python
help(math.factorial)
```

Output:
```
Help on built-in function factorial in module math:

factorial(x, /)
    Find x!.
    
    Raise a ValueError if x is negative or non-integral.
```

The `help` function can also be used to access the documentation for a module:

```python
help(math)
```

Output:
```
Help on module math:

NAME
    math

MODULE REFERENCE
    https://docs.python.org/3.9/library/math
    
    The following documentation is automatically generated from the Python
    source files.  It may be incomplete, incorrect or include features that
    are considered implementation detail and may vary between Python
    implementations.  When in doubt, consult the module reference at the
    location listed above.

DESCRIPTION
    This module provides access to the mathematical functions
    defined by the C standard.

FUNCTIONS
    acos(x, /)
        Return the arc cosine (measured in radians) of x.
        
        The result is between 0 and pi.
    
    acosh(x, /)
        Return the inverse hyperbolic cosine of x.
    
    asin(x, /)
        Return the arc sine (measured in radians) of x.
        
        The result is between -pi/2 and pi/2.
    
    asinh(x, /)
        Return the inverse hyperbolic sine of x.
    
    atan(x, /)
        Return the arc tangent (measured in 
```

In a notebook it is also possible to use `?` to access the documentation, which will be displayed at the bottom of the notebook:

```python
import random
random.gauss?
```

Output:
```
[0;31mSignature:[0m [0mrandom[0m[0;34m.[0m[0mgauss[0m[0;34m([0m[0mmu[0m[0;34m,[0m [0msigma[0m[0;34m)[0m[0;34m[0m[0;34m[0m[0m
[0;31mDocstring:[0m
Gaussian distribution.

mu is the mean, and sigma is the standard deviation.  This is
slightly faster than the normalvariate() function.

Not thread-safe without a lock around calls.
[0;31mFile:[0m      ~/.conda/envs/default/lib/python3.9/random.py
[0;31mType:[0m      method
```

## Exercise 05.1 (random numbers)

- Using the '`randint`' function from the '`random`' module (https://docs.python.org/3/library/random.html#random.randint), 
  develop a function `dice_roll` that emulates the roll of a dice with $n$ sides.

- For $n=6$, devise and implement a test to check that it is a fair dice. To solve this you could use a list to save the results of rolls. If you do not know how to do this then complete the next notebook before coming back to this part.

#### (a) Dice roll code:

<!-- grade_id: cell-70a83466894475ee, points: 0, grade: false, solution: true -->
```python
import random

help(random.randint)
```

Output:
```
Help on method randint in module random:

randint(a, b) method of random.Random instance
    Return random integer in range [a, b], including both end points.
```

```python
def dice_roll(n):
    a = random.randint(1, n) 
    return a

#test for n=6 
dice_roll(6)
```

Output:
```
3
```

<!-- grade_id: cell-0970743b0c035f78, points: 0, grade: true, solution: false -->
```python
for n in range(1, 20):
    for j in range(100):
        value = dice_roll(n) 
        assert value >= 1 and value <= n
```

#### (b) Test for fairness

<!-- grade_id: cell-0826760c62c21bff, points: 0, grade: true, solution: true -->
```python
# TEST 1 - NOT USING A LIST

#define number of sides on a dice 
n = 6

#define number of rolls
no_rolls = 100000

#varaible for how many ones were rolled
num_ones = 0

for i in range (no_rolls):
    a = dice_roll(n)
    if a == 1:
        num_ones = num_ones + 1

#if these numbers are similar then the test is far
print('No. of ones rolled: ', num_ones)
print('Expected frequency: ', no_rolls / 6)
```

Output:
```
No. of ones rolled:  16648
Expected frequency:  16666.666666666668
```

<!-- grade_id: cell-0826760c62c21bff, points: 0, grade: true, solution: true -->
```python
# TEST 2 - USING A LIST
import numpy as np

#define number of sides on a dice 
n = 6

#define number of rolls
no_rolls = 100000

#define list to record results in
results = []

#create a list of dice roll results 
for i in range (no_rolls):
    a = dice_roll(n)
    results.append(a)

#find the frequency of each dice roll
results_np = np.array(results)
np.unique(results_np, return_counts=True)
```

Output:
```
(array([1, 2, 3, 4, 5, 6]), array([16600, 16570, 16655, 16746, 16600, 16829]))
```

## Exercise 05.2 (data compression)

**This is an advanced exercises. You may not feel confident tackling this initially. You may want to come back and try this later.**

For devices with limited memory, data compression can be important. Data compression is
a field of its own, but with libraries we can compress (and uncompress) data easily without being expert in
the details.

Below is a program code for compressing a passage from Hamlet, by Shakespeare.

```python
# Import the compression module
import zlib

# Create a string that we wish to compress
text = """
Welcome, dear Rosencrantz and Guildenstern!
Moreover that we much did long to see you,
The need we have to use you did provoke
Our hasty sending. Something have you heard
Of Hamlet's transformation; so call it,
Sith nor the exterior nor the inward man
Resembles that it was. What it should be,
More than his father's death, that thus hath put him
So much from the understanding of himself,
I cannot dream of: I entreat you both,
That, being of so young days brought up with him,
And sith so neighbour'd to his youth and havior,
That you vouchsafe your rest here in our court
Some little time: so by your companies
To draw him on to pleasures, and to gather,
So much as from occasion you may glean,
Whether aught, to us unknown, afflicts him thus,
That, open'd, lies within our remedy."""

# Convert Python string to bytes and check type
text_bytes = text.encode("utf-8")
print(type(text_bytes))

# Get number of bytes used to store string
print("Number of bytes for uncompressed string:", len(text_bytes))

# Compress string and get number of byes used for compressed string
text_comp = zlib.compress(text_bytes)
print("Number of bytes for compressed string:", len(text_comp))

# Display the compression efficiency
print("Compression efficiency: ", len(text_comp)/len(text_bytes))

# Decompress the string
text_decomp = zlib.decompress(text_comp)

# Check that original and decompressed string are the same (more on aseret)
if text != text_decomp.decode("utf-8"):
    print("Problem: original and decompressed string differ.")
```

Output:
```
<class 'bytes'>
Number of bytes for uncompressed string: 785
Number of bytes for compressed string: 466
Compression efficiency:  0.5936305732484076
```

Using the above as a guide, examine the compression efficiency of 

1. Compressing one large string made up of the the passage by Shakespeare repeated 100 times; and
2. Compressing a random string of the same length as the repeated Shakespeare passage.

To help you, the below function generates random string of length `N`:

```python
import random
import string

def random_string(N):
    return ''.join([random.choice(string.ascii_letters + string.digits) for n in range(N)])

print(random_string(8))
```

Output:
```
asiF99xc
```

### Solution

```python
# Create a string
text = """
Welcome, dear Rosencrantz and Guildenstern!
Moreover that we much did long to see you,
The need we have to use you did provoke
Our hasty sending. Something have you heard
Of Hamlet's transformation; so call it,
Sith nor the exterior nor the inward man
Resembles that it was. What it should be,
More than his father's death, that thus hath put him
So much from the understanding of himself,
I cannot dream of: I entreat you both,
That, being of so young days brought up with him,
And sith so neighbour'd to his youth and havior,
That you vouchsafe your rest here in our court
Some little time: so by your companies
To draw him on to pleasures, and to gather,
So much as from occasion you may glean,
Whether aught, to us unknown, afflicts him thus,
That, open'd, lies within our remedy."""
```

Import the necessary modules:

```python
import random
import string
import zlib
```

Repeat the Shakespeare string 100 times, and compress:

<!-- grade_id: cell-c08bb020d9be6620, points: 0, grade: false, solution: true -->
```python
# Create string of Shakespeare passage repeated 100 times
# YOUR CODE HERE
long_text = text * 100
long_text_bytes = long_text.encode("utf-8")
long_text_comp = zlib.compress(long_text_bytes)
```

Create random string and compress:

<!-- grade_id: cell-c5d8c07589ec3cd1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
rand_text = random_string(len(long_text))
rand_text_bytes = rand_text.encode("utf-8")
rand_text_comp = zlib.compress(rand_text_bytes)
```

Compare compression efficiency:

<!-- grade_id: cell-a2dd98ee370e22a3, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
# Display the compression efficiency
print("Compression efficiency: ", len(long_text_comp)/len(long_text_bytes))
print("Compression efficiency: ", len(rand_text_comp)/len(rand_text_bytes))
```

Output:
```
Compression efficiency:  0.01178343949044586
Compression efficiency:  0.7520764331210191
```