# Introduction

A data structure is an 'object' in a program that holds a collection of data. A simple
data structure might be an 'array' that holds the components of a vector, or a list of
names. A more complicated data structure might represent a telephone directory, 
holding [name, telephone number] pairs.

Modern languages, like Python, provide a range of library (built-in) data structures.
These are well-tested and optimised; it is good practice to use library data structures to make programs
simpler, easier to read and more efficient. 

It is possible to develop your own data structures for special purposes. In this activity
we focus on selecting and using built-in data structures. The creation
of your own data structures is introduced in Activity 12.


### Objectives

- Use `list`, `tuple` and dictionary (`dict`) data structures
- Use iteratation to visit entries in a data structure 
- Learn to select the right data structure for an application

# Data structures

So far we have restricted our examples to simple built-in data types, e.g. `string`, `int`, and `float`. In practice, programs will use *data structures* to collect data together into useful packages. For example, rather than representing a vector `r` of length 3 using three floats `u`, `v` and `w`, we could represent 
it as a list of floats, `r = [u, v, w]`. Similarly, if we want to store the names of students in a laboratory group, rather than a string variable for each student, we could work with a list of names, e.g.:

```python
lab_group0 = ["Sarah", "John", "Joe", "Emily"]
lab_group1 = ["Roger", "Rachel", "Amer", "Caroline", "Colin"]
```

This is a much more powerful construction because we can perform operations on a list, such as checking its length (number of students in a lab group), sort the names in the list into alphabetical order, and add or remove names. We could even make a list-of-lists, e.g.

```python
lab_groups = [lab_group0, lab_group1]
```

to collect all the lab groups together into a *nested* list.

Data structures are particularly useful when passing data to functions. Say we want a function that prints the names of students in a given lab group. Rather than passing the name of each student (with different groups having differing numbers of members), we can pass just the list of lab group members to the function.
Similarly, we could develop a function that computes the dot product between two vectors of arbitrary length, passing to the function just the two vectors rather than each component.

We will look at three built-in Python data structures that are commonly used. They are:

- `list` 
- `tuple`
- `dict` (dictionary)

# Lists

A `list` is a sequence of data. An 'array' in most other languages is a similar concept, but Python lists are more general than most arrays as they can hold a mixture of types. A list is constructed using square brackets:

```python
lab_group0 = ["Sarah", "John", "Joe", "Emily", "Quang"]
print("Lab group members: {}".format(lab_group0))
print("Size of lab group: {}".format(len(lab_group0)))

print("Check the Python object type: {}".format(type(lab_group0)))
```

Output:
```
Lab group members: ['Sarah', 'John', 'Joe', 'Emily', 'Quang']
Size of lab group: 5
Check the Python object type: <class 'list'>
```

The function `len` returns the length (number of items) of the list.

An empty list is created by

```python
my_list = []
```

A list of length 5 with repeated values can be created by

```python
my_list = ["Oliver"]*5
print(my_list)
```

Output:
```
['Oliver', 'Oliver', 'Oliver', 'Oliver', 'Oliver']
```

## Iterating over lists

Looping over each item in a list (or more generally a sequence) is called *iterating*. We iterate over the members of the lab group using the syntax:

```python
for member in lab_group0:
    print(member)
```

Output:
```
Sarah
John
Joe
Emily
Quang
```

Say we want to iterate over the names of the lab group members, and get the position
of each member in the list. We use `enumerate` for this: 

```python
for n, member in enumerate(lab_group0):
    print(n, member)
```

Output:
```
0 Sarah
1 John
2 Joe
3 Emily
4 Quang
```

In the above, `n` is the position in the list and `member` is the $n$th entry in the list. Sometimes we need to know the position, in which case `enumerate` is helpful. However, when possible it is preferable to use 'plain' iteration. Note again that Python counts from zero - it uses zero-based indexing.

## Manipulating lists 

There are many functions for manipulating lists. It might be useful to sort the list:

```python
lab_group0.sort()
for member in lab_group0:
    print(member)
```

Output:
```
Emily
Joe
John
Quang
Sarah
```

In the above, `sort` is known as a 'method' of a `list`. It performs an *in-place* sort, i.e. `lab_group0` is sorted, rather than creating a new list with sorted entries (for the latter we would use `sorted(lab_group0)`, which returns a new list). More on methods later when we get to object-oriented design.

With lists we can add and remove students:

```python
# Remove the second student (indexing starts from 0, so 1 is the second element)
lab_group0.pop(1)
print(lab_group0)

# Add new student "Josephine" at the end of the list
lab_group0.append("Josephine")
print(lab_group0)
```

Output:
```
['Emily', 'John', 'Quang', 'Sarah']
['Emily', 'John', 'Quang', 'Sarah', 'Josephine']
```

or maybe join two groups to create one larger group: 

```python
lab_group1 = ["Roger", "Rachel", "Amer", "Caroline", "Colin"]

lab_group = lab_group0 + lab_group1
print(lab_group)
```

Output:
```
['Emily', 'John', 'Quang', 'Sarah', 'Josephine', 'Roger', 'Rachel', 'Amer', 'Caroline', 'Colin']
```

or create a list of group lists:

```python
lab_groups = [lab_group0, lab_group1]
print(lab_groups)

print("---")

print("Print each lab group (name and members):")
for i, lab_group in enumerate(lab_groups):
    print(i, lab_group)
```

Output:
```
[['Emily', 'John', 'Quang', 'Sarah', 'Josephine'], ['Roger', 'Rachel', 'Amer', 'Caroline', 'Colin']]
---
Print each lab group (name and members):
0 ['Emily', 'John', 'Quang', 'Sarah', 'Josephine']
1 ['Roger', 'Rachel', 'Amer', 'Caroline', 'Colin']
```

## Indexing

Lists store data in order, so it is possible to 'index into' a list using an integer (this will be familiar to anyone who has used C), e.g.:

```python
first_member = lab_group0[0]
third_member = lab_group0[2]
print(first_member, third_member)
```

Output:
```
Emily Quang
```

or

```python
for i in range(len(lab_group0)):
    print(lab_group0[i])
```

Output:
```
Emily
John
Quang
Sarah
Josephine
```

Indices start from zero, and run through to (length - 1).

Indexing can be useful for numerical computations. We use it here to compute the dot product of two vectors. A vector can be considered as a list of numbers. The dot product is when you take two lists of the same length, multiple togeter the numbers in the same position in each list, then sum the results.

```python
# Two vectors of length 4
x = [1.0, 3.5, 7.2, 8.9]
y = [-1.0, 27.1, 1.0, 6]

# Compute dot-product
dot_product = 0.0
for i in range(len(x)):
    dot_product += x[i]*y[i]

print(dot_product)
```

Output:
```
154.45000000000002
```

When possible, it is better to iterate over a list rather than use indexing, since there are data structures that support iterating but do not support indexing.

If we have a list-of-lists,

```python
lab_group0 = ["Sarah", "John", "Joe", "Emily"]
lab_group1 = ["Roger", "Rachel", "Amer", "Caroline", "Colin"]
lab_groups = [lab_group0, lab_group1]
```

we can use the first index to access a list, and a second index to access the entry in that list:

```python
group = lab_groups[0]
print(group)

name = lab_groups[1][2]
print(name)
```

Output:
```
['Sarah', 'John', 'Joe', 'Emily']
Amer
```

## Heterogeneity (lists with mixed types)

Python lists are heterogeneous data structures - this means they can store mixed types, e.g.

```python
mixed_list = ["Adam", 2 + 4j, 1.0, 4]
for entry in mixed_list:
    print(entry, type(entry))
```

Output:
```
Adam <class 'str'>
(2+4j) <class 'complex'>
1.0 <class 'float'>
4 <class 'int'>
```

Arrays in most languages are homogeneous - all types in the array must be the same.

There are *many* ways in which lists can be manipulated. The best way to learn how to perform a specific operation is to use a search engine.

## List comprehension

A powerful construct in modern languages, including Python, is *list comprehension*. It is a way to succinctly build lists from other lists. It can be very useful, but should be applied sensibly as it can sometimes be difficult to read. 

For more information on the format and approach to list comprehensions, you can see the following links:
- https://www.w3schools.com/python/python_lists_comprehension.asp
- https://realpython.com/list-comprehension-python/

Say we have a list of numbers and we wish to create a new list that squares each number in the original list and adds 5. Using list comprehension:

```python
x = [4, 6, 10, 11]
y = [a*a + 5 for a in x] 

print(f"Using list comprehension will return {x} and {y}")

#this is the same as
y = []

for a in x:
    y.append(a*a + 5)
    
print(f"Longer method will return {x} and {y}")
```

Output:
```
Using list comprehension will return [4, 6, 10, 11] and [21, 41, 105, 126]
Longer method will return [4, 6, 10, 11] and [21, 41, 105, 126]
```

To understand the meaning, read the statement left-to-right.

As another example, say we have a list of names and we want to 

- build a new list of names that contains only the names with more than 5 characters; and 
- for these names we want to add a full stop at the end.

Using list comprehension:

```python
lab_group1 = ["Roger", "Rachel", "Amer", "Caroline", "Colin"]
group = [name + "." for name in lab_group1 if len(name) > 5]

print(lab_group1)
print(group)
```

Output:
```
['Roger', 'Rachel', 'Amer', 'Caroline', 'Colin']
['Rachel.', 'Caroline.']
```

Mastering list comprehension is not simple, but it is very powerful.

## Tuples

Tuples are closely related to lists. The main difference is that a tuple cannot be changed after it has been created. 
In computing jargon, it is *immutable*.

For something that should not change after it has been created, such as a vector of length three with fixed entries, a tuple is more appropriate than a list. It is 'safer' in this case since it cannot be modified accidentally in a program. Being immutable ('read-only') also permits implementations to possibly exploit this to optimise for speed.

To create a tuple, use round brackets. Say at a college each student is assigned a room, and the rooms are numbered. A student 'Laura' is given room 32:

```python
room = ("Laura", 32)
print("Room allocation: {}".format(room))
print("Length of entry: {}".format(len(room)))
print(type(room))
```

Output:
```
Room allocation: ('Laura', 32)
Length of entry: 2
<class 'tuple'>
```

We can iterate over tuples in the same way as with lists,

```python
# Iterate over tuple values
for d in room:
    print(d)
```

Output:
```
Laura
32
```

and we can index into a tuple:

```python
print(room[1])
print(room[0])
```

Output:
```
32
Laura
```

We might have a student who is given permission to live outside of college. To keep track of them we still want an entry for the student, but there is no associated room number. We can create a tuple of length one using '`("a",)`', e.g.:

```python
room = ("Adrian",)
print("Room allocation: {}".format(room))
print("Length of entry: {}".format(len(room)))
print(type(room))
```

Output:
```
Room allocation: ('Adrian',)
Length of entry: 1
<class 'tuple'>
```

As part of a rooms database, we can create a list of tuples:

```python
room_allocation = [("Adrian",), ("Laura", 32), ("John", 31), ("Penelope", 28), ("Fraser", 28), ("Gaurav", 19)]
print(room_allocation)
```

Output:
```
[('Adrian',), ('Laura', 32), ('John', 31), ('Penelope', 28), ('Fraser', 28), ('Gaurav', 19)]
```

To make it easier for students to find their rooms on a printed room list, we can sort the list:

```python
room_allocation.sort()
print(room_allocation)
```

Output:
```
[('Adrian',), ('Fraser', 28), ('Gaurav', 19), ('John', 31), ('Laura', 32), ('Penelope', 28)]
```

We could improve the printed list by not printing those living outside of the college:

```python
for entry in room_allocation:
    if len(entry) > 1:
        print("Name: {} \n  Room: {}".format(entry[0], entry[1])) #\n creates new indented line
```

Output:
```
Name: Fraser 
  Room: 28
Name: Gaurav 
  Room: 19
Name: John 
  Room: 31
Name: Laura 
  Room: 32
Name: Penelope 
  Room: 28
```

In summary, prefer tuples over lists when the length will not change.

# Dictionaries (maps)

We used a list of tuples in the previous section to store room allocations. If we wanted to find which room a particular student has been allocated we would need to iterate through the list and check each name. For a very large list, this might not be very efficient.

There is a better way to do this, using a 'dictionary' (or sometimes called a 'map'). We have used indexing (with integers) into lists and tuples for direct access to a specific entry. This works if we know the index to the entry of interest. But, for a room list we identify individuals by name rather than a contiguous set of integers. Using a dictionary, we can build a 'map' from names (the *keys*) to room numbers (the *values*). 

A Python dictionary (`dict`) is declared using curly braces:

```python
room_allocation = {"Adrian": None, "Laura": 32, "John": 31, "Penelope": 28, "Fraser": 28, "Gaurav": 19}
print(room_allocation)
print(type(room_allocation))
```

Output:
```
{'Adrian': None, 'Laura': 32, 'John': 31, 'Penelope': 28, 'Fraser': 28, 'Gaurav': 19}
<class 'dict'>
```

Each entry is separated by a comma. For each entry we have a 'key', which is followed by a colon, and then the 'value'. Note that for Adrian we have used '`None`' for the value, which is a Python keyword for 'nothing' or 'empty'.

Now if we want to know which room Fraser has been allocated, we can query the dictionary by key:

```python
frasers_room = room_allocation["Fraser"]
print(frasers_room)
```

Output:
```
28
```

If we try to use a key that does not exist in the dictionary, e.g.

    frasers_room = room_allocation["Frasers"]

Python will give an error (raise an exception). If we're not sure that a key is present, we can check:

```python
print("Fraser" in room_allocation)
print("Frasers" in room_allocation)
```

Output:
```
True
False
```

(We can also use '`in`' to check if an entry exists in a list or tuple).
We can iterate over the keys in a dictionary:

```python
for d in room_allocation:
    print(d)
```

Output:
```
Adrian
Laura
John
Penelope
Fraser
Gaurav
```

or iterate over both the keys and the values:

```python
for name, room_number in room_allocation.items():
    print(name, room_number)
```

Output:
```
Adrian None
Laura 32
John 31
Penelope 28
Fraser 28
Gaurav 19
```

Note that the order of the printed entries in the dictionary is different from the input order. This is because a dictionary stores data differently from a list or tuple. Lists and tuples store entries 'linearly' in memory
(contiguous pieces of memory), which is why we can access entries by index. Dictionaries use a different type of storage which allows us to perform look-ups using a 'key'.

We have used a string as the key so far, which is common. However, we can use almost any type as a key, and we can mix types. For example, we might want to 'invert' the room allocation dictionary to create a room-to-name map: 

```python
# Create empty dictionary
room_allocation_inverse = {}

# Build inverse dictionary to map 'room number' -> name 
for name, room_number in room_allocation.items():
    # Insert entry into dictionary
    room_allocation_inverse[room_number] = name

print(room_allocation_inverse)
```

Output:
```
{None: 'Adrian', 32: 'Laura', 31: 'John', 28: 'Fraser', 19: 'Gaurav'}
```

We can now ask who is in room 28 and who is in room 29. Not all rooms are occupied, so we should include a check that the room number is a key in our dictionary:

```python
rooms_to_check = [28, 29]

for room in rooms_to_check:
    if room in room_allocation_inverse:
        print("Room {} is occupied by {}.".format(room, room_allocation_inverse[room]))
    else:
        print("Room {} is unoccupied.".format(room))
```

Output:
```
Room 28 is occupied by Fraser.
Room 29 is unoccupied.
```

# Choosing a data structure

An important task when developing a computer program is selecting the *appropriate* data structure for a task.
The more flexible the data structure, generally the less efficient it will be and it may use more memory compared to simpler data structures.

- If efficiency is not critical, pick a data structure that provides the functionality you need, and offers 
  flexibility and ease of use. Safety should also be considered (this can be a good reason for choosing a 
  tuple over a list).
- Use iterators rather than indexing when possible. This will allow switching from data structures that support 
  indexing to data structures that do not support indexing (such as a '`set`' data structure).

For numerical computations, efficiency is often essential and picking the right data structure is critical
for this. We will look at data structures that are specialied for numerical computations in the next notebook, and we'll see the huge difference in speed there can be between different data structures.

## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
#ORGINAL FUNCTION FROM EXERSICE 4
def area(x0, y0, x1, y1, x2, y2):
    """calculate area of a triangle based on its vertices cooridnates"""
    
    a = (x0*(y1 - y2) + x1*(y2 - y0) + x2*(y0 - y1))/2
    
    return abs(a)

#NEW FUNCTION - METHOD 1 have two lists containing x or y co-ordinates
# x = [x0, x1, x2]
# y = [y0, y1, y2]

def area_1(x, y):
    """calculate area of a triangle based on its vertices cooridnates"""

    #forms check
    assert len(x) == 3 and len(y) == 3
    
    #caclaute area 
    a = (x[0]*(y[1]-y[2]) + x[1]*(y[2]-y[0]) + x[2]*(y[0]-y[1]))/2
    
    return abs(a)

#NEW FUNCTION - METHOD 2 have 1 list containing turples of the verticies
# t = [(x0, y0), (x1, y1), (x2, y2)]

def area_2(t):
    """calculate area of a triangle based on its vertices cooridnates"""
    
    #check 3 vertices are inputted
    assert len(t) == 3   

    #check each vertices has 2 co-ordinates
    for vertex in t:
        assert len(vertex) == 2
    
    #define each set of verticies co-ordinates
    co_ord0 = t[0]
    co_ord1 = t[1]
    co_ord2 = t[2]
      
    #caclaute area    
    a = (co_ord0[0]*(co_ord1[1] - co_ord2[1]) +    #(x0*(y1 - y2) +
         co_ord1[0]*(co_ord2[1] - co_ord0[1]) +    #x1*(y2 - y0) +
         co_ord2[0]*(co_ord0[1] - co_ord1[1]) )/2  #x2*(y0 - y1))/2 
    
    return abs(a)
```

```python
#TEST THAT ALL 3 METHODS RETURN THE SAME AREA VALYE FOR THE TRIANGLE

#triangle verticies
x0, y0 = 0.0, 0.0
x1, y1 = 0.0, 2.0
x2, y2 = 3.0, 0.0

#orginal method
A = area(x0, y0, x1, y1, x2, y2)

#method 1
x = [x0, x1, x2]
y = [y0, y1, y2]

B = area_1(x, y)

#method 2
triangle_verticies = [(x0, y0), (x1, y1), (x2, y2)]

C = area_2(triangle_verticies)

print(A, B, C)
```

Output:
```
3.0 3.0 3.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon. Note that the sigma sign (big S-like symbol) indicates loop which sums over its contents, and the vertical bars around the formula tell you to take the absolute value of the thing inside it.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

```python
#example shape of a triangle which has 2 verticies
# shape = [(0,0), (0, 2), (3, 0)] 

def shape_area(shape):
    """Calculates area of a given polygon"""
    
    #define number of vertices present in the shape
    num_verticies = len(shape)
    
    #assign variable to where recorded area value will be stored
    area = 0
    
    #for each set of verticies we need to calculate the value of (Xi * Yi+1) - (Xi+1 * Yi)   
    for i in range(num_verticies-1):
        
        #create list of Xi, Yi and Xi+1, Yi+1 verticies
        vertex_i = shape[i]
        vertex_i_1 = shape[i+1]

        #evaluate the expression (Xi * Yi+1) - (Xi+1 * Yi)
        area += vertex_i[0]*vertex_i_1[1] - vertex_i_1[0]*vertex_i[1]  #incrementaly adds to find the sum of all the verticies 
    
    #find the absolute value of the sum of formula and multiply by 1/2
    area = abs(area*0.5) 
    
    return area

```

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
triangle = [(0,0), (0, 2), (3, 0)] 

triangle_area = shape_area(triangle)
print(triangle_area)

square = [(0,0), (0, 2), (2, 2), (2, 0)]

square_area = shape_area(square)
print(square_area)
```

Output:
```
3.0
4.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    
    #check vectors are the same size 
    if len(x) != len(y):
        print("Error - vectors x and y are not the same")
        
    #create new vector - where the sum of two other vectors will be stored
    z = []
    
    #using list indexing to add two vectors of arbitrary length
    for i in range(len(x)):
        
        new_element = x[i] + y[i]
        
        z.append(new_element)
        
    return z   

#method 2 - more memory effiecent to create list with zeros than appending 
def sum_vector_1(x, y):
    "Return sum of two vectors"
    if len(x) != len(y):
        print("Warning! vector sizes do not match")
    z = [0]*len(x)
    for i in range(len(x)):
        z[i] = x[i] + y[i]
    
    return z

```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
help(zip)
```

Output:
```
Help on class zip in module builtins:

class zip(object)
 |  zip(*iterables) --> A zip object yielding tuples until an input is exhausted.
 |  
 |     >>> list(zip('abcdefg', range(3), range(4)))
 |     [('a', 0, 0), ('b', 1, 1), ('c', 2, 2)]
 |  
 |  The zip object yields n-length tuples, where n is the number of iterables
 |  passed as positional arguments to zip().  The i-th element in every tuple
 |  comes from the i-th iterable argument to zip().  This continues until the
 |  shortest argument is exhausted.
 |  
 |  Methods defined here:
 |  
 |  __getattribute__(self, name, /)
 |      Return getattr(self, name).
 |  
 |  __iter__(self, /)
 |      Implement iter(self).
 |  
 |  __next__(self, /)
 |      Implement next(self).
 |  
 |  __reduce__(...)
 |      Return state information for pickling.
 |  
 |  ----------------------------------------------------------------------
 |  Static methods defined here:
 |  
 |  __new__(*args, **kwargs) from builtins.type
 |      Create and ret
```

```python
x = [0, 4.3, -5, 7]
y = [-2, 7, -15, 1]

#using function
z_1 = sum_vector(x, y)

#using list comprehension - zip iterates through multiple lists simultaneously
z_2 = [xi + yi for xi, yi in zip(x, y)]

print(z_1, z_2)
```

Output:
```
[-2, 11.3, -20, 8] [-2, 11.3, -20, 8]
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
colleges = {"Emmanuel": "EM",
            "Churchill": "CHU",
            "Darwin": "DAR",
            "Homerton": "HO",
            "Girton": "G"}

#create dictionary going from abbreviation to name
colleges_inverse = {}

for name, abbreviation in colleges.items():
    colleges_inverse[abbreviation] = name

print(colleges)
print("-----")
print(colleges_inverse)

#create list of college abbreviations sorted into alphabetical order
ordered_college_abbr = sorted(list(colleges.values()))

print(ordered_college_abbr)


```

Output:
```
{'Emmanuel': 'EM', 'Churchill': 'CHU', 'Darwin': 'DAR', 'Homerton': 'HO', 'Girton': 'G'}
-----
{'EM': 'Emmanuel', 'CHU': 'Churchill', 'DAR': 'Darwin', 'HO': 'Homerton', 'G': 'Girton'}
['CHU', 'DAR', 'EM', 'G', 'HO']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```